// variables1.rs
// Make me compile! Execute the command `rustlings hint variables1` if you want a hint :)

// 让其编译通过！

// About this `I AM NOT DONE` thing:
// We sometimes encourage you to keep trying things on a given exercise,
// even after you already figured it out. If you got everything working and
// feel ready for the next exercise, remove the `I AM NOT DONE` comment below.

//

fn main() {
    let x = 5; // 思考：let 有什么意义
    println!("x has the value {}", x);
}
